# LSTM para series temporales

**Laboratorio:** `05_lstm_time_series`  
**Dataset real:** seoul_bike  
**Fuente:** UCI  
**Versión:** estudiante

## Objetivo

Pronosticar demanda horaria respetando el orden temporal.

**Pregunta guía:** ¿Cómo cambia el error cuando aumenta el horizonte y cómo evita el protocolo mirar el futuro?

## Contrato científico

1. `split_seed` define las particiones.  
2. `training_seed` define inicialización y batches.  
3. Toda selección ocurre con validation.  
4. `experiment.lock.json` congela decisiones antes de abrir test.  
5. Test se informa una sola vez.

```python
from pathlib import Path
import json
import yaml
import torch

from neural_labs.catalog import ROOT, get_lab, get_dataset
from neural_labs.datasets import prepare_dataset, audit_bundle
from neural_labs.baselines import run_baseline
from neural_labs.experiments import run_lab

LAB_ID = '05_lstm_time_series'
SPLIT_SEED = 42
TRAINING_SEED = 42
QUICK = True
```

```python
lab = get_lab(LAB_ID)
dataset_card = get_dataset(LAB_ID)
lab, dataset_card
```

## Procedencia, licencia y limitaciones

Lee la ficha antes de descargar. No redistribuyas datos si la licencia no lo permite.

```python
manifest_path = ROOT / 'labs' / LAB_ID / 'data' / 'dataset.yaml'
manifest = yaml.safe_load(manifest_path.read_text(encoding='utf-8'))
manifest
```

## Preparación y auditoría

La siguiente celda descarga datos reales. Puede requerir Internet, credenciales de Kaggle o aceptación de condiciones.

```python
# bundle = prepare_dataset(LAB_ID, quick=QUICK, seed=SPLIT_SEED)
# audit_bundle(bundle)
print('Descomenta para descargar y preparar el dataset real.')
```

## Exploración específica del dominio: time_series

Realiza backtesting walk-forward, evalúa por horizonte y compara LSTM, TCN, persistencia y media móvil.

```python
# Comprobación del orden cronológico y ventanas
# print(bundle.summary)
# assert bundle.summary.get('chronological_split', True)
```

## Línea base sobre validation

La línea base no debe mirar test antes de congelar el experimento.

```python
# baseline_validation = run_baseline(LAB_ID, bundle, quick=QUICK, evaluation_split='validation')
# baseline_validation
```

## Modelo y teoría

**Arquitectura:** `lstm_regression`  
**Fundamento:** Puertas input, forget y output de una LSTM.

```python
# sample, target = bundle.train[0]
# print('input:', sample.shape, 'target:', target)
# print('selection metric:', lab['selection_metric'])
```

## Entrenamiento reproducible

El comando guarda contrato de inferencia, lock experimental, métricas, checkpoint y tarjetas.

```python
# result = run_lab(
#     LAB_ID, quick=QUICK, config_name='baseline',
#     split_seed=SPLIT_SEED, training_seed=TRAINING_SEED, device='auto',
# )
# result.run_dir, result.metrics
```

## Inspección de artefactos

```python
# sorted(path.name for path in result.run_dir.iterdir())
```

## Ejercicio evaluable

<!-- grade_id: experiment_design, points: 5, grade: false, solution: true -->
```python
# YOUR CODE HERE
EXPERIMENT_DESCRIPTION = ''
EXPECTED_SEEDS = []
```

<!-- grade_id: experiment_design_test, points: 5, grade: true, solution: false -->
```python
assert isinstance(EXPERIMENT_DESCRIPTION, str) and len(EXPERIMENT_DESCRIPTION) >= 20
assert len(EXPECTED_SEEDS) >= 3
```

## Interpretación y responsabilidad

Documenta resultados negativos, sesgos, grupos con peor desempeño, costo computacional y usos no recomendados.

<!-- grade_id: conclusion, points: 5, grade: false, solution: true -->
```python
# YOUR CODE HERE
CONCLUSION = ''
```

<!-- grade_id: conclusion_test, points: 5, grade: true, solution: false -->
```python
assert len(CONCLUSION) >= 40
```

## Próximos pasos

Ejecuta primero con `--quick`; después usa el dataset completo, varias semillas y el pipeline de benchmark.